# SLU01 - Pandas 101: Exercise notebook

```python
# First things first, import:
import pandas as pd
import numpy as np
import hashlib
```

- In this notebook the following is tested:

    - Creating pandas Series
    - Creating pandas DataFrame
    - Load a dataset from a file
    - Preview a dataframe
    - Convert datatypes in DataFrame

## Exercise 1: Series

### 1.1 Create a pandas Series from a list.

Create a `list` with the names of the items given in the cell below. Then create a `Series` using that list, using the order provided.

Note: Let pandas generate the index on its own. Also, don't forget to delete the raise error!

<!-- grade_id: cell-d0be99db1c88ce43, points: 0, grade: false, solution: true -->
```python
# items (each of type string): pork, milk, peas, peanuts, potatoes
# items_list = 

# Create a series with the items above, and call it items_series
# items_series = 

# YOUR CODE HERE
items_list= ['pork', 'milk', 'peas', 'peanuts', 'potatoes']
items_series=pd.Series(items_list)
```

```python
# run this cell to see the results 
items_series
```

Output:
```
0        pork
1        milk
2        peas
3     peanuts
4    potatoes
dtype: object
```

<!-- grade_id: cell-0344c841b5dac488, points: 2, grade: true, solution: false -->
```python
assert(isinstance(items_list, list)), "The data structure used is not the correct one."
assert(len([i for i in items_list if not isinstance(i, str)]) == 0), "The items inside the list should be of type string."
assert(len(items_series) == 5), "The length of the list is not correct."
assert(hashlib.sha256(''.join(items_list).encode()).hexdigest() == '2b0898564acfeeacc4c2f2ae3fb7e8057552d46849f2523d5d4edb3a234528a2'), "The items in the list don't seem to be correct."
assert(hashlib.sha256(''.join(items_series.values.tolist()).encode()).hexdigest() == '2b0898564acfeeacc4c2f2ae3fb7e8057552d46849f2523d5d4edb3a234528a2'), "The items in the Series don't seem to be correct."
assert(np.equal(items_series.index.to_numpy(), np.arange(0, 5)).sum() == 5), "The Series index doesn't seem to be correct."
```

### 1.2 CO<sub>2</sub> equivalents

Let's now use the `items_list` as index and build a `Series` containing the kg CO<sub>2</sub> equivalents per item (the amount of CO<sub>2</sub> corresponding to the warming effect of greenhouse gases emitted during production of the item, [ref](https://doi.org/10.1126/science.aaq0216)).

```python
# RUN THIS CELL FIRST
# kg CO2 equivalents of the foods in items_list:
food_CO2 = np.array([7.6, 3.2, 0.4, 1.2, 0.6]) 
```

<!-- grade_id: cell-2d2283c13fa0600e, points: 0, grade: false, solution: true -->
```python
# Create a series which has items_list (created in exercise 1.1) as index, 
# and the respective value in the food_CO2 list given in the cell above as values.
# CO2_eq = 

# YOUR CODE HERE
CO2_eq=pd.Series(food_CO2, index=items_list)
```

```python
# run this cell to see the results 
CO2_eq
```

Output:
```
pork        7.6
milk        3.2
peas        0.4
peanuts     1.2
potatoes    0.6
dtype: float64
```

<!-- grade_id: cell-096e084aeb31ba69, points: 2, grade: true, solution: false -->
```python
assert (isinstance(CO2_eq, pd.Series)), "CO2_eq is not a series."
assert(len(CO2_eq==5)), "The length of CO2_eq is not correct."
assert(CO2_eq.sum() == 13), 'Something is wrong with the values of the series. Did you use the correct input?'
assert(hashlib.sha256(''.join(CO2_eq.index.tolist()).encode()).hexdigest() == '2b0898564acfeeacc4c2f2ae3fb7e8057552d46849f2523d5d4edb3a234528a2'), 'The order of the index is not correct.'
```

Cool! In 1 exercise you learned 2 things:
- that a pandas Series can be built using several data structures, such as lists or numpy arrays
- that if you want to reduce your carbon footprint and help save the planet, you should turn to a plant-based diet (<300 g meat and <1.7 kg dairy per week, see here: https://takethejump.org/)
ðŸ˜ƒ ðŸŒ

## Exercise 2: DataFrames

The goal of this exercise is to create a simple DataFrame from several data structures.

```python
# RUN THIS CELL FIRST
# this is the data you'll use to fill each column of your dataframe
github_repos = ['freeCodeCamp', 'free-programming-books', 'awesome', '996.ICU', 'coding-interview-university']
repo_creators = np.array(['freeCodeCamp', 'EbookFoundation', 'sindresorhus', '996icu', 'jwasham'])
number_of_forks = [33599, 57194, 26485, 21497, 69434]
number_of_stars = [374074, 298393, 269997, 267901, 265161]
```

<!-- grade_id: cell-b5c9767e3556cd9b, points: 0, grade: false, solution: true -->
```python
# Add the data from the lists github_repos, repo_creators, number_of_forks, number_of_stars
# to a dictionary called dict_most_popular_repos:
#   - use the 4 variables created in the cell above to fill the data for each key
#   - each key should be a string containing the name of the corresponding variable.
# dict_most_popular_repos = 

# Create a dataframe called df_most_popular_repos
#   - set an index with the values 'first', 'second', 'third', 'fourth', 'fifth'
#   - use the dictionary created above to populate the dataframe.
# df_most_popular_repos = ...


# YOUR CODE HERE
dict_most_popular_repos={'github_repos': github_repos, 
'repo_creators':repo_creators, 
'number_of_forks': number_of_forks, 
'number_of_stars': number_of_stars}

df_most_popular_repos=pd.DataFrame(dict_most_popular_repos, index= ('first', 'second', 'third', 'fourth', 'fifth'))
df_most_popular_repos
```

Output:
```
github_repos    repo_creators  number_of_forks  \
first                  freeCodeCamp     freeCodeCamp            33599   
second       free-programming-books  EbookFoundation            57194   
third                       awesome     sindresorhus            26485   
fourth                      996.ICU           996icu            21497   
fifth   coding-interview-university          jwasham            69434   

        number_of_stars  
first            374074  
second           298393  
third            269997  
fourth           267901  
fifth            265161
```

<!-- grade_id: cell-a4ab78569f17b22b, points: 2, grade: true, solution: false -->
```python
assert(list(dict_most_popular_repos.keys())==['github_repos', 'repo_creators', 'number_of_forks', 'number_of_stars']), "The dictionary keys are not correct."
assert(len(dict_most_popular_repos) == 4), "The length of the dictionary is not correct."
assert(isinstance(dict_most_popular_repos,dict)), 'dict_most_popular_repos is not a dictionary.'
assert(isinstance(df_most_popular_repos, pd.DataFrame)), 'df_most_popular_repos is not a DataFrame'
assert(df_most_popular_repos['github_repos'].tolist()==github_repos), "The 'github_repos' column doesn't look right."
assert(df_most_popular_repos['repo_creators'].tolist()==list(repo_creators)), "The 'repo_creators' column doesn't look right."
assert(df_most_popular_repos['number_of_forks'].tolist()==number_of_forks), "The 'number_of_forks' column doesn't look right."
assert(df_most_popular_repos['number_of_stars'].tolist()==number_of_stars), "The 'number_of_stars' column doesn't look right."
assert(df_most_popular_repos.shape == (5, 4)), 'The size of the dataframe is not correct.'
assert(df_most_popular_repos.index.tolist() != ('first', 'second', 'third', 'fourth','fifth')), 'The index is not correct. Reread the instructions.'
```

## Exercise 3: Loading DataFrames from files

### 3.1 Load a dataset into a dataframe
Let's load a dataset with data about chocolate bars. It is a subset from a Kaggle dataset available [here](https://www.kaggle.com/datasets/rtatman/chocolate-bar-ratings/).

<!-- grade_id: cell-6f209833953ab9cf, points: 0, grade: false, solution: true -->
```python
# Load the dataset from the file located at data/chocolate_bars.csv.
# ds_chocolate = 

# YOUR CODE HERE
ds_chocolate=pd.read_csv('chocolate_bars.csv')
```

```python
# Print the dataframe head() to get an idea of what you've just loaded.
ds_chocolate.head()
```

Output:
```
co                     name  date  pct  free       loc  \
0     Artisan du Chocolat                Venezuela  2010  100  True      U.K.   
1                  Bonnat              One Hundred  2006  100  True    France   
2  Bouga Cacao (Tulicorp)  El Oro, Hacienda de Oro  2009  100  True   Ecuador   
3                 C-Amaro                  Ecuador  2013  100  True     Italy   
4         Claudio Corallo                 Principe  2008  100  True  Sao Tome   

    rtg                 btp                  org  
0  1.75                   Â             Venezuela  
1  1.50                   Â                     Â   
2  1.50  Forastero (Arriba)              Ecuador  
3  3.50                   Â               Ecuador  
4  1.00           Forastero  Sao Tome & Principe
```

<!-- grade_id: cell-a08b45d8b2814c8f, points: 1, grade: true, solution: false -->
```python
assert(isinstance(ds_chocolate, pd.DataFrame)), "Something is wrong. ds_chocolate is not a dataframe."
assert(ds_chocolate.shape == (45, 9)), "The shape is not correct. Did you follow all the instructions in the comments?"
assert(sum(ds_chocolate.columns == ['co', 'name', 'date', 'pct', 'free', 'loc', 'rtg', 'btp', 'org']) == 9), "The columns don't look right."
assert(ds_chocolate.rtg.sum()+ds_chocolate.date.sum()+ds_chocolate.pct.sum() == 94886),"The content of the dataframe is not correct"
assert(hashlib.sha256(''.join(ds_chocolate.name.tolist()).encode()).hexdigest() == '9c1403a3605e81d77ee5379fc733705bace8235c4a5cf36f27f9af27440b68f3'), "The content of the dataframe is not correct."
```

### 3.2 Load a dataset, but this time better

Notice that the column names in the ds_chocolate dataframe are not very informative. This is not very useful to someone looking at the data. Instead we want to load the dataset with the following `column names` (in this order):
- `'company'` - maker of the chocolate bar
- `'bar_name'` - original name of the chocolate bar
- `'review_date'` - date when the rating was given
- `'cocoa_percentage'`
- `'sugarfree'` - contains sugar or not
- `'company_location'` 
- `'rating'` - from 1 to 5
- `'cocoa_bean_type'` - cocoa variety
- `'cocoa_bean_origin'` - where the bean type originated

<!-- grade_id: cell-8e15e2392133ead5, points: 0, grade: false, solution: true -->
```python
# Load the file at 'data/chocolate_bars.csv' into a dataframe ds_chocolate.
# set the column names to 
#  in this order.
# You will need to check the documentation at 
# https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.read_csv.html to see how you can do this.
# ds_chocolate = 


# YOUR CODE HERE
ds_chocolate=pd.read_csv('chocolate_bars.csv')
ds_chocolate=ds_chocolate.rename(columns={'co':'company', 'name':'bar_name', 'date':'review_date', 'pct':'cocoa_percentage', 'free':'sugarfree', 'loc':'company_location', 'rtg':'rating', 'btp':'cocoa_bean_type', 'org':'cocoa_bean_origin'})
ds_chocolate.to_csv('chocolate_bars.csv', index=False)
```

```python
# Print the dataframe head() to get an idea of what you've just loaded.
ds_chocolate.head()
```

Output:
```
company                 bar_name  review_date  \
0     Artisan du Chocolat                Venezuela         2010   
1                  Bonnat              One Hundred         2006   
2  Bouga Cacao (Tulicorp)  El Oro, Hacienda de Oro         2009   
3                 C-Amaro                  Ecuador         2013   
4         Claudio Corallo                 Principe         2008   

   cocoa_percentage  sugarfree company_location  rating     cocoa_bean_type  \
0               100       True             U.K.    1.75                   Â    
1               100       True           France    1.50                   Â    
2               100       True          Ecuador    1.50  Forastero (Arriba)   
3               100       True            Italy    3.50                   Â    
4               100       True         Sao Tome    1.00           Forastero   

     cocoa_bean_origin  
0            Venezuela  
1                    Â   
2              Ecuador  
3              Ecuador  
4  Sao Tome
```

<!-- grade_id: cell-4457e5154efddfdf, points: 2, grade: true, solution: false -->
```python
assert(isinstance(ds_chocolate, pd.DataFrame)), "Something is wrong. ds_chocolate is not a dataframe."
assert(ds_chocolate.shape == (45, 9)), "The shape is not correct. Did you follow all the instructions in the comments?"
assert(sum(ds_chocolate.columns == ['company','bar_name','review_date','cocoa_percentage','sugarfree','company_location','rating','cocoa_bean_type','cocoa_bean_origin']) == 9), "The columns don't look right."
assert(ds_chocolate.rating.sum()+ds_chocolate.review_date.sum()+ds_chocolate.cocoa_percentage.sum() == 94886),"The content of the dataframe is not correct"
assert(hashlib.sha256(''.join(ds_chocolate.bar_name.tolist()).encode()).hexdigest() == '9c1403a3605e81d77ee5379fc733705bace8235c4a5cf36f27f9af27440b68f3'), "The content of the dataframe is not correct."
assert(hashlib.sha256(''.join(ds_chocolate.columns).encode()).hexdigest() == '9681020039c95c14566f818837a704bafa96b014f0438921179a1ff52e2286ed'), "The column names are not correct"
```

### 3.3 Preview the datatypes

<!-- grade_id: cell-e9e62f9d35948d34, points: 0, grade: false, solution: true -->
```python
# Store the datatypes of all columns of ds_chocolate in ds_chocolate_dtypes.
# Use the method you learned in the learning notebook.
# ds_chocolate_dtypes = 

# Note: if you used the correct method, 
# the result will be a pandas series containing the datatypes of each column,
# with the index formed by the column names


# YOUR CODE HERE
ds_chocolate_dtypes=ds_chocolate.dtypes
```

```python
# Check your output - there should be object, float, bool, and integer types.
ds_chocolate_dtypes
```

Output:
```
company               object
bar_name              object
review_date            int64
cocoa_percentage       int64
sugarfree               bool
company_location      object
rating               float64
cocoa_bean_type       object
cocoa_bean_origin     object
dtype: object
```

<!-- grade_id: cell-efab0f02db3e97d7, points: 2, grade: true, solution: false -->
```python
assert(sum([x in ds_chocolate_dtypes.index for x in ds_chocolate.columns]) == 9), "The index of ds_chocolate_dtypes should contain all columns in ds_chocolate."
assert(hashlib.sha256(''.join([str(typ) for typ in ds_chocolate_dtypes]).encode()).hexdigest() == '10f3499fbbb3d16acebe010d689e3d64246a7cd6dff7bdb18cbfa845d5caf3c0'), "The dtypes are not correct."
```

### 3.4 Set the correct datatypes
The datatypes in `ds_chocolate` were infered, so all `strings` are set as `objects`. Convert all these `object` datatypes to `string` using a function you learned in the learning notebook.

<!-- grade_id: cell-383fa34cad0538db, points: 0, grade: false, solution: true -->
```python
# Set the correct datatypes in the ds_chocolate dataframe - convert the objects to strings.
# Store the new dtypes in the variable ds_chocolate_dtypes_converted.
# ds_chocolate = 
# ds_chocolate_dtypes_converted = 


# YOUR CODE HERE
ds_chocolate=ds_chocolate.convert_dtypes()
ds_chocolate_dtypes_converted=ds_chocolate.dtypes
```

```python
# Check you solution and compare it to the result of exercise 3.3. There will be pandas datatypes now.
ds_chocolate_dtypes_converted
```

Output:
```
company              string[python]
bar_name             string[python]
review_date                   Int64
cocoa_percentage              Int64
sugarfree                   boolean
company_location     string[python]
rating                      Float64
cocoa_bean_type      string[python]
cocoa_bean_origin    string[python]
dtype: object
```

<!-- grade_id: cell-bbb1770c9e8e0521, points: 1, grade: true, solution: false -->
```python
assert(sum([x in ds_chocolate_dtypes_converted.index for x in ds_chocolate.columns]) == 9), "The index of ds_chocolate_dtypes_converted should contain all columns in ds_jobs."
assert(hashlib.sha256(str(ds_chocolate_dtypes_converted['company']).encode()).hexdigest() == '473287f8298dba7163a897908958f7c0eae733e25d2e027992ea2edc9bed2fa8'), "The dtype of column 'company' is not as expected."
assert(hashlib.sha256(str(ds_chocolate_dtypes_converted['cocoa_percentage']).lower().encode()).hexdigest() == 'bc2229666b96007e875c5f62897ee5b7648db2baa5fabf3e771afac323afbd57'), "The dtype of column 'cocoa_percentage' is not as expected."
assert(hashlib.sha256(str(ds_chocolate_dtypes_converted['sugarfree'])[:4].encode()).hexdigest()=='b760f44fa5965c2474a3b471467a22c43185152129295af588b022ae50b50903'), "The dtype of column 'sugarfree' is not as expected."
```

### 3.5 Get information about the dataframe size
Use a method you learned in the learning notebook to retrieve the `number of rows` and the `number of columns` in the ds_chocolate dataframe.

<!-- grade_id: cell-e4bf7a0d20b4d075, points: 0, grade: false, solution: true -->
```python
# number_of_rows = 
# number_of_columns = 


# YOUR CODE HERE
ds_chocolate.shape
number_of_rows, number_of_columns = ds_chocolate.shape
print(f'number of rows: {number_of_rows}')
print(f'number of column: {number_of_columns}')
```

Output:
```
number of rows: 45
number of column: 9
```

<!-- grade_id: cell-a6082129b0c73871, points: 2, grade: true, solution: false -->
```python
assert(hashlib.sha256(str(int(number_of_rows)).encode()).hexdigest() == '811786ad1ae74adfdd20dd0372abaaebc6246e343aebd01da0bfc4c02bf0106c'), "The number of rows is not correct."
assert(hashlib.sha256(str(int(number_of_columns)).encode()).hexdigest() == '19581e27de7ced00ff1ce50b2047e7a567c76b1cbaebabe5ef03f7c3017bb5b7'), "The number of columns is not correct."
```

### 3.6 Load a json file into a dataframe
Let's load a new dataframe called hdi from the file stored at `data/HDI.json`. It's the human development index statistics in the years 1990-2019, a subset of a kaggle dataset available [here](https://www.kaggle.com/datasets/elmartini/human-development-index-historical-data).

<!-- grade_id: cell-23ac4f8fe231e90c, points: 0, grade: false, solution: true -->
```python
# Load the datafile from data/HDI.json and store it in the variable hdi. Use the appropriate method for json files.
# hdi = 


# YOUR CODE HERE
hdi=pd.read_json('HDI.json')
```

```python
# Preview your dataframe
hdi.head()
```

Output:
```
HDI Rank      Country   1991   1996   2001   2006   2011   2016   2019
0       169  Afghanistan  0.307  0.335  0.353  0.429  0.477  0.502  0.511
1        69      Albania  0.631  0.646  0.678  0.713  0.764  0.788  0.795
2        91      Algeria  0.576  0.602  0.647  0.690  0.728  0.743  0.748
3        36      Andorra    NaN    NaN  0.815  0.837  0.836  0.866  0.868
4       148       Angola    NaN    NaN  0.410  0.473  0.533  0.578  0.581
```

<!-- grade_id: cell-71ea15bcdf2d8056, points: 1, grade: true, solution: false -->
```python
assert(isinstance(hdi, pd.DataFrame)), "hdi is not a dataframe."
assert(hdi.shape == (189, 9)), "The shape of the dataframe is not correct.?"
assert(sum(hdi.columns == ['HDI Rank', 'Country', '1991', '1996', '2001', '2006', '2011', '2016', '2019']) == 9), "The columns don't look right."
assert(hashlib.sha256(''.join(hdi.Country.tolist()).encode()).hexdigest() == '8471ec328ba910fd961813c7619ae8776684a396672a959975015b24b3335ca3'), "The Country column looks wrong."
assert(hdi['HDI Rank'].sum() == 17914), "Something is wrong with the 'HDI Rank' column."
assert(hdi['2001'].sum() == 111.082), "Something is wrong with the '2019' column."
```

### 3.7 Get a numpy array of column names
Store the names of the `columns` in the hdi dataframe as a `numpy array`.

<!-- grade_id: cell-4fd5c93627335cd6, points: 0, grade: false, solution: true -->
```python
# First extract the columns into hdi_cols.
# hdi_cols = 

# Then convert the output into a NumPy array.
# hdi_cols_array = 


# YOUR CODE HERE
hdi_cols=hdi.columns
hdi_cols_array=hdi_cols.to_numpy()
```

```python
# Always preview your variables to see the result of the operations.
print(hdi_cols, type(hdi_cols), "\n", sep="\n")
print(hdi_cols_array, type(hdi_cols_array), sep="\n")
```

Output:
```
Index(['HDI Rank', 'Country', '1991', '1996', '2001', '2006', '2011', '2016',
       '2019'],
      dtype='object')
<class 'pandas.core.indexes.base.Index'>


['HDI Rank' 'Country' '1991' '1996' '2001' '2006' '2011' '2016' '2019']
<class 'numpy.ndarray'>
```

<!-- grade_id: cell-4dce9e138896815d, points: 2, grade: true, solution: false -->
```python
assert(isinstance(hdi_cols, pd.core.indexes.base.Index)), "Use the method you learned to extract the columns into hdi_cols."
assert(len(hdi_cols) == 9), "There are 9 columns in the hdi dataframe. Did you extract them all? Also, make sure you don't change the variable hdi."
assert(isinstance(hdi_cols_array, np.ndarray)), "The hdi_cols_array is not a numpy array."
assert(hashlib.sha256(''.join(hdi_cols).encode()).hexdigest() == 'd29520e57d10ac733aad2b8a0e5b044083fc150a5110543e05f5a0c82699cd8f'),"The extracted column names are not correct."
```

### 3.8 Extract the index as a numpy array
Do the same as in exercise 3.7, but now for the index of hdi.

<!-- grade_id: cell-3681531b79d28fd3, points: 0, grade: false, solution: true -->
```python
# Extract the index using the method you learned.
# hdi_index = 

# Convert it to a numpy array.
# hdi_index_array = 


# YOUR CODE HERE
hdi_index=hdi.index
hdi_index_array=hdi.index.to_numpy()
hdi_index_array
```

Output:
```
array([  0,   1,   2,   3,   4,   5,   6,   7,   8,   9,  10,  11,  12,
        13,  14,  15,  16,  17,  18,  19,  20,  21,  22,  23,  24,  25,
        26,  27,  28,  29,  30,  31,  32,  33,  34,  35,  36,  37,  38,
        39,  40,  41,  42,  43,  44,  45,  46,  47,  48,  49,  50,  51,
        52,  53,  54,  55,  56,  57,  58,  59,  60,  61,  62,  63,  64,
        65,  66,  67,  68,  69,  70,  71,  72,  73,  74,  75,  76,  77,
        78,  79,  80,  81,  82,  83,  84,  85,  86,  87,  88,  89,  90,
        91,  92,  93,  94,  95,  96,  97,  98,  99, 100, 101, 102, 103,
       104, 105, 106, 107, 108, 109, 110, 111, 112, 113, 114, 115, 116,
       117, 118, 119, 120, 121, 122, 123, 124, 125, 126, 127, 128, 129,
       130, 131, 132, 133, 134, 135, 136, 137, 138, 139, 140, 141, 142,
       143, 144, 145, 146, 147, 148, 149, 150, 151, 152, 153, 154, 155,
       156, 157, 158, 159, 160, 161, 162, 163, 164, 165, 166, 167, 168,
       169, 170, 171, 172, 173, 174, 175, 176, 177, 178, 179, 18
```

<!-- grade_id: cell-1bf51b9a2de225bc, points: 1, grade: true, solution: false -->
```python
assert(isinstance(hdi_index, pd.core.indexes.base.Index)), "Use the method you learned to extract the index into hdi_index."
assert(len(hdi_index) == 189), "The length of the hdi_index variable is incorrect."
assert(sum(hdi_index_array) == 17766), "Something is wrong with the index array."
assert(isinstance(hdi_index_array, np.ndarray)), "The hdi_index_array does not look like a numpy array."
```


### 3.9 Describe the data in your dataframe
Last but not least, remember how you can get some stats and info on your dataframe? If you don't, make sure to reread the learning notebook. If you do, let's jump to this final exercise.

Using only the two methods you learned to get information and statistics on a dataframe answer the three questions in the cell below manually.

```python
# Use this draft cell to print stuff to help you answer the questions below.
hdi.describe().T
```

Output:
```
count       mean        std    min       25%      50%        75%  \
HDI Rank  189.0  94.783069  54.754486  1.000  48.00000  95.0000  142.00000   
1991      144.0   0.601854   0.166086  0.220   0.48075   0.6340    0.73300   
1996      148.0   0.621007   0.168904  0.244   0.47650   0.6465    0.75700   
2001      174.0   0.638402   0.169386  0.268   0.48150   0.6670    0.77700   
2006      186.0   0.666468   0.163872  0.300   0.51950   0.6985    0.79550   
2011      188.0   0.692713   0.155098  0.338   0.56350   0.7240    0.80600   
2016      188.0   0.713367   0.151803  0.378   0.58800   0.7380    0.82625   
2019      189.0   0.722423   0.149791  0.394   0.60200   0.7400    0.82900   

              max  
HDI Rank  189.000  
1991        0.871  
1996        0.890  
2001        0.914  
2006        0.934  
2011        0.942  
2016        0.950  
2019        0.957
```

<!-- grade_id: cell-3fd00abe2d329bde, points: 0, grade: false, solution: true -->
```python
# Question 1
# What is the mean value for HDI in the year 2011 (rounded to 2 decimal points)?
# mean_HDI_2011 = 

# Question 2
# What is the maximum value for HDI in the year 1996 (round to 2 decimal points)?
# max_HDI_1996 = 

# Question 3 
# How many non-null entries do we have for the year 1991? Store the answer as an integer.
# nonnull_HDI_2001 = 



# YOUR CODE HERE
mean_HDI_2011=hdi['2011'].mean()
mean_HDI_2011_rounded=round(mean_HDI_2011, 2)
print(f'The mean value for HDI in the year 2011 is {mean_HDI_2011_rounded}')

max_HDI_1996=hdi['1996'].max()
max_HDI_1996_rounded=round(max_HDI_1996, 2)
print(f'The maximum value for HDI in the year 1996 is {max_HDI_1996_rounded}')

nonnull_HDI_2001=int(hdi['2001'].notnull().sum())
print(f'The year 1991 has {nonnull_HDI_2001} non-null entries')
```

Output:
```
The mean value for HDI in the year 2011 is 0.69
The maximum value for HDI in the year 1996 is 0.89
The year 1991 has 174 non-null entries
```

<!-- grade_id: cell-ea52883fc83bdbfb, points: 2, grade: true, solution: false -->
```python
assert (isinstance(mean_HDI_2011, float)), "mean_HDI_2011 should be a float."
assert (isinstance(max_HDI_1996, float)), "max_HDI_1996 should be a float."
assert (isinstance(nonnull_HDI_2001, int)), "nonnull_HDI_2001 should be an integer."
assert(hashlib.sha256(str(round(mean_HDI_2011,2)).encode()).hexdigest() == '18db98e3760ea46503aa9f79ad850cc408d0192446a5666220ca3b99346e3e8c'), "mean_HDI_2011 does not look right."
assert(hashlib.sha256(str(round(max_HDI_1996,2)).encode()).hexdigest() == '0003dc8ba4e379d078ccc278b8c13ca665b2440770df16f54511bf7609203239'), "max_HDI_1996 does not look right."
assert(hashlib.sha256(str(int(nonnull_HDI_2001)).encode()).hexdigest() == '41e521adf8ae7a0f419ee06e1d9fb794162369237b46f64bf5b2b9969b0bcd2e'), "nonnull_HDI_2001 does not look right."
```