```python
import hashlib # for grading
import json

# Standard imports
import numpy as np
import pandas as pd
import re
import string
import math

# NLTK imports
import nltk
nltk.download('stopwords')
from nltk.tokenize import WordPunctTokenizer
from nltk.stem.snowball import SnowballStemmer
from nltk.corpus import stopwords

# SKLearn related imports
from sklearn.feature_extraction.text import CountVectorizer
from sklearn.feature_extraction.text import TfidfTransformer
from sklearn.pipeline import Pipeline
from sklearn.base import TransformerMixin, BaseEstimator

from sklearn.linear_model import LogisticRegression
from sklearn.model_selection import train_test_split
from sklearn.metrics import classification_report
from sklearn.metrics import accuracy_score
```

## Exercise 1 - Price of cars (regex)
For the first question, you will be using regex to extract information from the `cars.txt` dataset. In this dataset, you'll find a list of cars that have been sold, as well as their brand, model and selling price.

Start by loading the data into a list. The list items are the lines of the file.

```python
path = "data/cars.txt"
cars = []
with open(path, 'r', encoding='utf-8') as f:
    cars = [l.strip() for l in f.readlines()]
```

```python
cars[:10]
```

In the first item, for example, `FORD` is the brand name, `Focus` is the model, and `19757` is the price.

### Exercise 1.1 - Find Toyota cars

First, we want to see which `TOYOTA` models have been sold. Find the items in the `cars` list which correspond to cars of the `TOYOTA` brand. Put these items into a list called `ans_1_1`.

<!-- grade_id: cell-9d7292940dd0f0b2, points: 0, grade: false, solution: true -->
```python
# ans_1_1 = ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9d5ce92d78b9b815, points: 0.5, grade: true, solution: false -->
```python
assert isinstance(ans_1_1,list)
assert len(ans_1_1) == 111
assert hashlib.sha256(json.dumps(''.join(sorted(ans_1_1))).encode()).hexdigest() == \
'292f64bf78d500d2bb5dc13ca282edb571f1589b47ce0353f030de70c630608c', 'Not correct, try again.'
```

### Exercise 1.2 - Find models with numbers

Next, find the items in the `cars` list whose model is a set of numbers instead of characters. For example, `'BMW -- 535 -- 23521'`. Store these items in the list called `ans_1_2`.

<!-- grade_id: cell-8dccb1f349ce02a4, points: 0, grade: false, solution: true -->
```python
# ans_1_2 = ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-98c96bad0127da20, points: 0.5, grade: true, solution: false -->
```python
assert isinstance(ans_1_2,list)
assert len(ans_1_2) == 73
assert hashlib.sha256(json.dumps(''.join(sorted(ans_1_2))).encode()).hexdigest() == \
'cdee14b87d092c0510a1cb9af05e8b9e96fc6a1c81489f7fd235e4a24e6b9119', 'Not correct, try again.'
```

### Exercise 1.3 - Selling price bellow 1000

Finally, get the car brands and models whose selling price is below 1000.

Save the results in the list `ans_1_3`. Each element in this list should be in the format `BRAND -- MODEL`.

<!-- grade_id: cell-b8146ce9f0ba50df, points: 0, grade: false, solution: true -->
```python
# ans_1_3 = ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b746057a51e92a5d, points: 0.5, grade: true, solution: false -->
```python
assert len(ans_1_3) == 77
assert hashlib.sha256(json.dumps(''.join(sorted(ans_1_3))).encode()).hexdigest() == \
'4d1ef7339260bc7b24f98c0ea6031b4a6fff0f7a9e0100dfde1e17dbbe1308f9', 'Not correct, try again.'
```

## Exercise 2 - Job postings
The challenge of this exercise notebook is to classify job postings as 'Fake' or 'Real'. In this exercise, we'll be preprocessing the data.

Let's load the data:

```python
df = pd.read_csv('data/job_postings.csv', index_col=0).convert_dtypes()

X = df['description']
y = df['fraudulent']
```

```python
df.head()
```

Let's look at an example of a job description and its corresponding label:

```python
df.iloc[10]['description']
```

```python
df.iloc[10]['fraudulent']
```

Let's check the data size and distribution of classes:

```python
def get_data_stats(X, y):
    print(f"Size of dataset: {len(X)}")
    unique, counts = np.unique(y, return_counts=True)
    print(f"Distribution of classes: {dict(zip(unique, counts))}")

get_data_stats(X, y)
```

The classes are evenly distributed. We'll use a dev and test set to be able to identify overfitting and check the performance on unseen data.

**Note**: So far you've used the `train`/`val`/`test` nomenclature for naming variables related to training, validation and test sets, respectively. `dev` is short for "development" and is just another typical identifier for the validation set, and we'll use it throughout this notebook instead of `val`.

```python
# train dev test split
X_train, X_temp, y_train, y_temp = train_test_split(X, y, test_size=0.2, random_state=42, stratify=y)
X_dev, X_test, y_dev, y_test = train_test_split(X_temp, y_temp, test_size=0.3, random_state=42, stratify=y_temp)
print(f"Train size: {len(X_train)}\nDev size: {len(X_dev)}\nTest size: {len(X_test)}")
```

The first step in the workflow is preprocessing which we'll do in this exercise.

### Exercise 2.1 - Tokenization

Implement the function `apply_tokenizer`. The function should receive a pandas series of text data like `X_train` and an NLTK-style tokenizer. It should return the series with the tokenized text. The tokens of each item in the series should be joined into one string with spaces in between like in the example in the docstring.

<!-- grade_id: cell-a8b7930a61806e32, points: 0, grade: false, solution: true -->
```python
def apply_tokenizer(data, tokenizer):
    """
    Tokenizes text data in the provided series using the provided tokenizer.
    
    E.g. for text data  "This is a test! No, it can't be"
         it returns "This is a test ! No , it can ' t be"
    
    Args:
    data - pd.Series with text data
    tokenizer - nltk tokenizer
    """
    
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-f829c5a222c54690, points: 1, grade: true, solution: false -->
```python
tokenizer = WordPunctTokenizer()
X_train_tok = apply_tokenizer(X_train, tokenizer)

assert isinstance(X_train_tok, pd.Series), 'The function should return a pandas series.'
assert len(X_train_tok) == 1530, 'The length of the series is not correct.'
assert isinstance(X_train_tok.iloc[0],str), 'The items of the series should be strings.'
assert hashlib.sha256(json.dumps(''.join([i for i in X_train_tok])).encode()).hexdigest() == \
'0e4a3d7fdab35e43079953d4d1328450b8b36454083f28966e3026e813eb3e3f', 'Not correct, try again.'
```

### Exercise 2.2 - Lowercasing

In the second step, implement a function that will lowercase the data.  It should take and return a pandas series.

<!-- grade_id: cell-ee47fb5a45fbd622, points: 0, grade: false, solution: true -->
```python
def apply_lowercase(data):
    """
    Lowercases the text data in the provided pandas series.
    
    Args:
    data - pd.Series
    """

    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-7979e12840663ea2, points: 1, grade: true, solution: false -->
```python
X_train_tok_lc = apply_lowercase(X_train_tok)

assert isinstance(X_train_tok_lc, pd.Series), 'The output should be a pandas series.'
assert len(X_train_tok_lc) == 1530, 'The length of the output is not correct.'
assert isinstance(X_train_tok_lc.iloc[0],str), 'The items of the series should be strings.'
assert hashlib.sha256(json.dumps(''.join([i for i in X_train_tok_lc])).encode()).hexdigest() == \
'bd90f00143ed79470e7c4000aa5a800a0c0f68fe0e661eb5ee126fc0d2bb38ac', 'Not correct, try again.'
```

### Exercise 2.3 - Stopwords

Now implement a function that filters the stopwords from the text data. The function should take and return a pandas series. We will use NLTK's built-in English stopword list shown below.

```python
stopword_list = stopwords.words('english')
```

<!-- grade_id: cell-42ecb29c8fe117f1, points: 0, grade: false, solution: true -->
```python
def apply_filter_stopwords(data, stopword_list):
    """
    Removes stopwords from the provided pandas series with text data.
    
    Args:
    data - pd.Series
    stopword_list - list of stopwords to filter out
    """
    
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-f70da0255ea6e291, points: 1, grade: true, solution: false -->
```python
X_train_tok_lc_nosw = apply_filter_stopwords(X_train_tok_lc, stopword_list)

assert isinstance (X_train_tok_lc_nosw, pd.Series), 'The output should be a pandas series.'
assert len(X_train_tok_lc_nosw) == 1530, 'The length of the output is not correct.'
assert isinstance(X_train_tok_lc_nosw.iloc[0],str), 'The items of the series should be strings.'
assert hashlib.sha256(json.dumps(''.join([i for i in X_train_tok_lc_nosw])).encode()).hexdigest() == \
'13a80f1f4d7740b20cd3e2993ab73d37563212a8b208d3f554646a5783234e87', 'Not correct, try again.'
```

### Exercise 2.4 - Punctuation

After filtering the stopwords, we want to remove punctuation from the text. Consider only the punctuation characters in `string.punctuation`. Make sure to remove all punctuation and not only tokens that are single punctuation characters. 

<!-- grade_id: cell-a0cd3cf5cc97a8f2, points: 0, grade: false, solution: true -->
```python
def apply_filter_punct(data):
    """
    Removes punctuation from the provided pandas series with text data.
    
    Args:
    data - pandas series with text data
    """

    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
X_train_tok_lc_nosw_nopunct = apply_filter_punct(X_train_tok_lc_nosw)
```

**Normalize whitespaces**

Run the following function on `X_train_tok_lc_nosw_nopunct` before checking your answers to remove extra white spaces.

```python
def normalize_whitespace(text):
    return re.sub(r"^\s+|\s+$|(?<=\s)\s*", "", text)

X_train_tok_lc_nosw_nopunct_norm = X_train_tok_lc_nosw_nopunct.apply(normalize_whitespace)
```

<!-- grade_id: cell-650a677a3a01d1bf, points: 1, grade: true, solution: false -->
```python
assert isinstance (X_train_tok_lc_nosw_nopunct_norm, pd.Series), 'The output should be a pandas series.'
assert len(X_train_tok_lc_nosw_nopunct_norm) == 1530, 'The length of the output is not correct.'
assert isinstance(X_train_tok_lc_nosw_nopunct_norm.iloc[0],str), 'The items of the series should be strings.'
assert hashlib.sha256(json.dumps(''.join([i for i in X_train_tok_lc_nosw_nopunct_norm])).encode()).hexdigest() == \
'a8273b9bae9dfb421af75eacc22bfd1b9c809ff36526060d756ffcbbed35dce9', 'Not correct, try again.'
```

### Exercise 2.5 - Stemming

The last preprocessing step that you are going to implement is stemming. Implement the function below to receive an NLTK-style stemmer and a pandas series with text data and return the series with the stemmed text.

<!-- grade_id: cell-a831ba989f3e50e6, points: 0, grade: false, solution: true -->
```python
def apply_stemmer(data, stemmer):
    """
    Stems the text data in the provided pandas series.
    
    Args:
    data - pd.Series
    stemmer - NLTK-style stemmer
    """

    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-3596a6510ebbda3d, points: 1, grade: true, solution: false -->
```python
stemmer = SnowballStemmer("english")
X_train_tok_lc_nosw_nopunct_norm_stem = apply_stemmer(X_train_tok_lc_nosw_nopunct_norm, stemmer)

assert isinstance (X_train_tok_lc_nosw_nopunct_norm_stem, pd.Series), 'The output should be a pandas series.'
assert len(X_train_tok_lc_nosw_nopunct_norm_stem) == 1530, 'The length of the output is not correct.'
assert isinstance(X_train_tok_lc_nosw_nopunct_norm_stem.iloc[0],str), 'The items of the series should be strings.'
assert hashlib.sha256(json.dumps(''.join([i for i in X_train_tok_lc_nosw_nopunct_norm_stem])).encode()).hexdigest() == \
'fb150190fd6957d636e470fb703db7a353e4d53fcdff4ef131431dac6461a895', 'Not correct, try again.'
```

### Exercise 2.6 - Everything together

Finally, join all the preprocessing steps from above into a transformer that applies the steps in the following order:
* tokenization
* lowercasing
* filtering stopwords
* filtering punctuation
* normalizing whitespace
* stemming.

Make use of the functions you designed above and don't forget to initialize all the necessary parameters.

<!-- grade_id: cell-ea5b2305431c20dd, points: 0, grade: false, solution: true -->
```python
class TextCleanerTransformer(TransformerMixin, BaseEstimator):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
text_cleaner = TextCleanerTransformer(
    WordPunctTokenizer(),
    stopwords=stopwords.words('english'),
    stemmer=SnowballStemmer("english"),
)

X_train_pre = text_cleaner.fit_transform(X_train)
```

<!-- grade_id: cell-4a87d0c9b1f20f7e, points: 2.5, grade: true, solution: false -->
```python
assert isinstance(X_train_pre, pd.Series), 'The output should be a pandas series.'
assert len(X_train_pre) == 1530, 'The length of the output is not correct.'
assert isinstance(X_train_pre.iloc[0],str), 'The items of the series should be strings.'
assert hashlib.sha256(json.dumps(''.join([i for i in X_train_pre])).encode()).hexdigest() == \
'fb150190fd6957d636e470fb703db7a353e4d53fcdff4ef131431dac6461a895', 'Not correct, try again.'
```

## Exercise 3 - Text classification

We will now classify the job postings as fake or real. Let's first load the preprocessed data and check the balance of the classes.

We are loading the preprocessed csv file here. This way you won't be penalized if you did not finish exercise 2.

```python
def load_dataset(file_name):
    """
    Loads a csv file and returns two pandas series, one
    the text and one containing the labels
    
    Args:
    file_name: path to input file
    """
    df = pd.read_csv(file_name, index_col = 0)

    return df['description'], df['fraudulent']
```

```python
X_train_pre, y_train = load_dataset('data/job_postings_train_preprocessed.csv')
X_dev_pre, y_dev = load_dataset('data/job_postings_dev_preprocessed.csv')
X_test_pre, y_test = load_dataset('data/job_postings_test_preprocessed.csv')
```

```python
get_data_stats(X_train_pre, y_train)
```

```python
get_data_stats(X_dev_pre, y_dev)
```

So we should be aiming for much better than 45% accuracy, which is what we would get if we naively predicted `1` (fake) for everything.

### Exercise 3.1 - Top N-grams in the BoW vectorization

First, we'll look at the top X N-grams in each category to see if anything is interesting. Implement a function that returns the most common N-grams and their count for the given label in the dataset. Use the `CountVectorizer` to create the N-grams. The function should return a list of tuples of the form `(N-gram, count)`, sorted by the `count` in descending order.

<!-- grade_id: cell-36567f8f392a4836, points: 0, grade: false, solution: true -->
```python
def top_ngrams_for_category(data, labels, filter_label, top_n=10, ngram_size=1):
    """
    Finds the top_n N-grams in the BoW for the given class label.
    
    Args:
    data: pd.Series with text data
    labels: class labels for the data
    filter_label: the label to filter the data on before getting ngrams
    top_n: top n N-grams to return
    ngram_size: the "N" in N-gram (e.g. if ngram_size=2, return only bigrams)
    
    Returns: list of tuples (N-gram, count)
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-cca0190cfba4ad83, points: 1, grade: true, solution: false -->
```python
top_10_unigrams_real = top_ngrams_for_category(X_train_pre, y_train, 0, top_n=10, ngram_size=1)
assert top_10_unigrams_real == [('nan', 3199),
                                ('work', 2689),
                                ('experi', 1966),
                                ('manag', 1854),
                                ('team', 1830),
                                ('servic', 1773),
                                ('develop', 1706),
                                ('custom', 1624),
                                ('compani', 1430),
                                ('time', 1428)]
top_6_unigrams_fake = top_ngrams_for_category(X_train_pre, y_train, 1, top_n=6, ngram_size=1)
assert top_6_unigrams_fake == [('nan', 3318),
                               ('work', 1828),
                               ('manag', 1303),
                               ('experi', 1287),
                               ('time', 1221),
                               ('servic', 1187)]
top_5_bigrams_real = top_ngrams_for_category(X_train_pre, y_train, 0, top_n=5, ngram_size=2)
assert top_5_bigrams_real == [('nan nan', 1349),
                              ('full time', 745),
                              ('custom servic', 398),
                              ('bachelor degre', 308),
                              ('year experi', 211)]
top_10_trigrams_fake = top_ngrams_for_category(X_train_pre, y_train, 1, top_n=10, ngram_size=3)
assert top_10_trigrams_fake == [('nan nan nan', 851),
                                ('nan full time', 165),
                                ('time nan nan', 138),
                                ('high school equival', 134),
                                ('full time nan', 131),
                                ('oil gas industri', 123),
                                ('time entri level', 116),
                                ('full time entri', 104),
                                ('level high school', 92),
                                ('mid senior level', 89)]
```

Looking at the top ngrams for each category, it doesn't seem like a BoW model will be very interesting, but let's try anyway.

### Exercise 3.2 - Modeling pipeline on BoW
Let's streamline our pipeline in a nice function. The function should set up the pipeline, fit the pipeline with the train data and predict on the dev data. It should also calculate the accuracy of the prediction and print the classification report.

The pipeline should have two steps, vectorization with sklearn `CountVectorizer` and classification with the `LogisticRegression`. Name the pipeline steps `vect` and `clf`. The `CountVectorizer` should take the given `ngram_range` and `max_features` parameter values.

The function should return the fitted pipeline, the prediction, and the accuracy of the prediction.

<!-- grade_id: cell-af90809e6c74e780, points: 0, grade: false, solution: true -->
```python
def train_and_validate(X_train, X_dev, y_train, y_dev, ngram_range=(1,1), max_features=None):
    """
    Train a model using sklearn's pipeline and return it along with the predictions and the
    accuracy in the validation set. Print the classification report as well.
    
    Args:
    X_train - preprocessed training data
    X_dev - preprocessed dev data
    y_train - labels of training data
    y_dev - labels of dev data
    ngram_range - ngram range to use in CountVectorizer (tuple)
    max_features - max number of features to use in CountVectorizer (int)
    
    Returns:
    text_clf - fitted pipeline
    y_dev_pred - prediction on the dev data (np.array)
    acc - accuracy of the prediction (float)
    """
    
    # YOUR CODE HERE
    raise NotImplementedError()
    print(classification_report(y_dev, y_dev_pred))
    return text_clf, y_dev_pred, acc
```

<!-- grade_id: cell-97ba1cceed5a53f9, points: 2, grade: true, solution: false -->
```python
clf, y_dev_pred, acc = train_and_validate(X_train_pre, X_dev_pre, y_train, y_dev)
assert isinstance(clf['vect'], CountVectorizer), 'The pipeline steps are not correct.'
assert isinstance(clf['clf'], LogisticRegression), 'The pipeline steps are not correct.'
assert hashlib.sha256(json.dumps(''.join([str(i) for i in y_dev_pred])).encode()).hexdigest() == \
'36720a1e3d54d5a2d740cfe8cff6370341459fb18ecfbbddc04e5978ea0e7a7e', 'The prediction is not correct.'
np.testing.assert_almost_equal(acc, 0.906, decimal=3, err_msg="The accuracy is not correct.")
```

We should look at some misclassified examples:

```python
for text, pred, true in zip(X_dev_pre[:70], y_dev_pred[:70], y_dev[:70]):
    if pred != true:
        print(f"Job posting: {text}")
        print(f"Predicted: {pred}, Actual: {true}\n")
```

So just with the simplest BoW model we already get an accuracy of 0.9! But let's see if we can do even better... In the misclassified examples, the last one even contains an application url but classified as 'fake'. That's suspicious...

### Exercise 3.3 - Tune hyperparameters
Run the function from the previous exercise for different N-gram ranges and/or with different values for max_features. Try to achieve an accuracy higher or equal to 0.9.

<!-- grade_id: cell-e75ac8652cc7e741, points: 0, grade: false, solution: true -->
```python
# clf_tuned, y_dev_pred_tuned, acc_tuned = train_and_validate(...)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-722d1b4ff466bdd7, points: 1, grade: true, solution: false -->
```python
assert(acc_tuned >= 0.90)
print(acc_tuned)
```

Now evaluate your model on the test set!

```python
y_test_pred = clf.predict(X_test)
print(classification_report(y_test, y_test_pred))
```

Great! We were able to improve the model a little on the dev set and the performance on the test set is pretty good!

## Exercise 4 - TF-IDF

We will now work with the TF-IDF vectorization of the preprocessed job postings data.

### Exercise 4.1 - Manual TF-IDF

First, implement a function that manually calculates the TF-IDF vectorization from a dataframe with a Bag of Words vectorization. Here is a reminder of the TF-IDF formula:

$$ tf\text{-}idf(t,d) = tf(t,d) \times idf(t) = \frac{n_{td}}{n_d} \times (log{(\frac{n+1}{df_t+1})} + 1 )$$

where $t$ is the term (word) for which we are calculating the weight, $d$ is the document containing the term, $n_{td}$ is the word count of the term $t$ in document $d$, $n_d$ is the number of words in document $d$, $n$ is the number of documents, and $df_t$ is the number of documents where the term appears.

```python
# We'll start you off with the BoW representation, in pandas dataframe format
vec = CountVectorizer()
bow_train = vec.fit_transform(X_train_pre)
bow_train_df = pd.DataFrame(bow_train.todense())
vocab = vec.vocabulary_
```

<!-- grade_id: cell-0b915bab713aba13, points: 0, grade: false, solution: true -->
```python
def calculate_tfidf(bow_df):
    """
    Calculates the tfidf vectorization from a BoW vectorization.

    Args:
    bow_df - dataframe with document word counts (Bag of Words)

    Returns:
    tfidf - dataframe with the calculated tfidf
    """
    # YOUR CODE HERE
    raise NotImplementedError()
    return tf_idf
```

<!-- grade_id: cell-86090d8867691e1a, points: 2, grade: true, solution: false -->
```python
tfidf_df = calculate_tfidf(bow_train_df)
assert tfidf_df.shape==bow_train_df.shape, 'The shape of the tfidf dataframe is not correct.'
assert ((tfidf_df>0).sum()==(bow_train_df>0).sum()).sum()==bow_train_df.shape[1], ''
assert hashlib.sha256(json.dumps(''.join([str(round(i,3)) for i in tfidf_df[(tfidf_df>0)].sum().to_numpy()])).encode()).hexdigest() \
== 'cdeee674f9f9449ffb7a1cec047e75a0318af7836d0b94fe3948ce81e1ae4559', 'Not correct, try again.'
```

### Exercise 4.2 - Top words in the TF-IDF vectorization

Implement a function which returns the top N most important words for the given class in the TF-IDF vectorization, i.e. the words with the highest sum of weights. The function takes the dataframe with the TF-IDF weights calculated in the previous exercise and a vocabulary of words and feature indices from the CountVectorizer and returns a list of the top N most important words.

<!-- grade_id: cell-00f14b79ff68d32e, points: 0, grade: false, solution: true -->
```python
def top_words_for_category_tfidf(data, labels, filter_label, vocab, top_n=10):
    """
    Finds the top_n words in the TF-IDF for the given class label.
    
    Args:
    data: pd.Series with text data
    labels: class labels for the data
    filter_label: the label to filter the data on before getting ngrams
    vocab: vocabulary of words and feature indices
    top_n: top n ngrams to return
    
    Returns: list of top-n words
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-3b51198b450763c2, points: 2, grade: true, solution: false -->
```python
top_15_fake = top_words_for_category_tfidf(tfidf_df, y_train, 1, vocab, top_n=15)
assert top_15_fake == ['nan',
                       'work',
                       'time',
                       'servic',
                       'entri',
                       'custom',
                       'posit',
                       'skill',
                       'amp',
                       'experi',
                       'home',
                       'manag',
                       'requir',
                       'administr',
                       'data']
top_20_real = top_words_for_category_tfidf(tfidf_df, y_train, 0, vocab, top_n=20)
assert top_20_real == ['nan',
                       'work',
                       'develop',
                       'manag',
                       'experi',
                       'custom',
                       'sale',
                       'team',
                       'servic',
                       'product',
                       'job',
                       'market',
                       'client',
                       'busi',
                       'design',
                       'compani',
                       'time',
                       'technolog',
                       'us',
                       'engin']
```

### Exercise 4.3 - Modeling pipeline on TF-IDF
Let's include the TF-IDF features into the pipeline. The function should set up the pipeline, fit the pipeline with the train data and predict on the dev data. It should also calculate the accuracy of the prediction and print the classification report.

The pipeline should have three steps: vectorization with sklearn `CountVectorizer`, transformation with sklearn `TfidfTransformer`, and classification with the `LogisticRegression`. Name the pipeline steps `vect`, `tfidf`, and `clf`. The `CountVectorizer` should take the given `ngram_range`, `max_features`, `max_df`, and `min_df` parameter values.

The function should return the fitted pipeline, the prediction, and the accuracy of the prediction.

<!-- grade_id: cell-02453bb681cd33b8, points: 0, grade: false, solution: true -->
```python
def train_and_validate_with_tfidf(X_train, X_dev, y_train, y_dev, ngram_range=(1,1),
                                  max_features=None, max_df=1.0, min_df=1):
    """
    Train a model using sklearn's pipeline and return it along with the predictions and the
    accuracy in the validation set. Print the classification report as well.
    
    Args:
    X_train - preprocessed training data
    X_dev - preprocessed dev data
    y_train - labels of training data
    y_dev - labels of dev data
    ngram_range - ngram range to use in CountVectorizer (tuple)
    max_features - max number of features to use in CountVectorizer (int)
    max_df - minimum threshold for document frequency to use in CountVectorizer (int or float)
    min_df - maximum threshold for document frequency to use in CountVectorizer (int or float)
    
    Returns:
    text_clf - fitted pipeline
    y_dev_pred - prediction on the dev data (np.array)
    acc - accuracy of the prediction (float)
    """
    
    # YOUR CODE HERE
    raise NotImplementedError()
    print(classification_report(y_dev, y_dev_pred))
    return text_clf, y_dev_pred, acc
```

<!-- grade_id: cell-ee1c189fe642928a, points: 2, grade: true, solution: false -->
```python
clf_tfidf, y_dev_pred_tfidf, acc_tfidf = train_and_validate_with_tfidf(X_train_pre, X_dev_pre, y_train, y_dev)
assert isinstance(clf_tfidf['vect'], CountVectorizer), 'The pipeline steps are not correct.'
assert isinstance(clf_tfidf['tfidf'], TfidfTransformer), 'The pipeline steps are not correct.'
assert isinstance(clf_tfidf['clf'], LogisticRegression), 'The pipeline steps are not correct.'
assert hashlib.sha256(json.dumps(''.join([str(i) for i in y_dev_pred_tfidf])).encode()).hexdigest() == \
'475ecac8a2c6e2dea9aa06f0105f19aa20ead066339a6398af8d797bfeafbd0a', 'The prediction is not correct.'
np.testing.assert_almost_equal(acc_tfidf, 0.906, decimal=3, err_msg="The accuracy is not correct.")
```

As before, we should also look at some misclassified examples:

```python
for text, pred, true in zip(X_dev_pre[:70], y_dev_pred_tfidf[:70], y_dev[:70]):
    if pred != true:
        print(f"Job posting: {text}")
        print(f"Predicted: {pred}, Actual: {true}\n")
```

Unfortunately, we're still not able to correctly classify those job postings. Let's keep going and see if we can do even better.

### Exercise 4.4 - Tune hyperparameters again

Use the `train_and_validate_with_tfidf` function you created before to train with different hyperparameters and get an accuracy score above 92% on the validation dataset. (This threshold is the same as what we got for the plain CountVectorizer pipeline).

<!-- grade_id: cell-07b6f3694941ac8, points: 0, grade: false, solution: true -->
```python
# clf_tfidf_tuned, y_dev_pred_tfidf_tuned, acc_tfidf_tuned = train_and_validate_with_tfidf(...)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-07b6f3694941ac81, points: 1, grade: true, solution: false -->
```python
assert(acc_tfidf_tuned > 0.92)
print(acc_tfidf_tuned)
```

Now evaluate your model on the test set!

```python
y_test_pred = clf_tfidf_tuned.predict(X_test)
print(classification_report(y_test, y_test_pred))
```

Did the performace on the test set improve against the pure BoW pipeline? If not, the pipeline might require more tuning.

Congratulations to mastering the first NLP unit! This was no easy task, be proud of yourself!