# BLU12 - Workflow - Exercise notebook

```python
import hashlib
import json
```

The main idea of the workflow unit is for you to understand how you should build a strategy for designing a recommender system.
For this purpose we recommend that you:

1. Review the workflow notebook carefully and make sure you are able to explain the whole pipeline to a 5-year-old child (or a [rubber duck](https://en.wikipedia.org/wiki/Rubber_duck_debugging)). Even better, try to get some ideas on how to improve it!

2. Choose one of these [datasets](https://github.com/caserec/Datasets-for-Recommender-Systems) and follow the steps in the Simplified Pipeline notebook.

3. Design and implement the whole RS pipeline for another dataset from stratch - as will happen in the Hackathon.

When you feel ready for the Hackathon, complete the next cell and submit it to the portal.

<!-- grade_id: cell-f38e2faccb3fe9d8, points: 0, grade: false, solution: true -->
```python
def ds_recommenders_oath_of_office():
    """The Recommender Systems' Data Scientist Oath of Office
    
    Returns:
        oath (str): should contain the following sentence:
                    'I do solemnly swear I have reviewed the Workflow notebook and I am ready for the Hackathon!'
    """
    # YOUR CODE HERE
    raise NotImplementedError()

oath = ds_recommenders_oath_of_office()
```

<!-- grade_id: cell-7867e029322ef452, points: 20, grade: true, solution: false -->
```python
assert hashlib.sha256(json.dumps(oath).encode()).hexdigest() == \
'e54e1d36b37cae24224bcb9f133f0f47735239c86349d7ad8ae5f965e8d90cd6'
```